---
title: Test
description: Test
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: d543d724ac
lang: python
xp: 100
skills: 2
```

This is the assignment text. It should help provide students with the background information needed.
The instructions that follow should be in bullet point form with clear guidance for what is expected.

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{python}
# Load datasets and packages here.
```

`@sample_code`
```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: cdf306881f
xp: 50
```

`@projector_key`
9bf853a62e1f87c8c531e29724f96369

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 29be126ec0
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 387e477c8a
xp: 50
```



`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: TabExercise
key: 480ea87ff2
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 56c664402a
xp: 35
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 48983d1d62
xp: 30
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 6116ba5abe
xp: 35
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: ea04996ef0
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: f4242b06a5
```

`@question`
lknlk

`@possible_answers`


`@hint`


`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: BulletExercise
key: 6d2a307fa9
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 193d9291e2
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: c50ac2a493
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
