---
title: 'chapter 2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: d40561d1b1
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 023c52a7a1
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: TabExercise
key: d324c41b08
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

***

```yaml
type: NormalExercise
key: b3bb0a33f7
xp: 20
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 318b99f7c4
xp: 20
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 52b08dbc0c
xp: 20
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: c03fcbe34d
xp: 20
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 5a575f6310
xp: 20
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: e32badf329
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: TabExercise
key: 27dde049e3
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 52e1eb2a32
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: b2e2a6a0ea
xp: 50
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```
